Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

# Current through a resistor

```python
import numpy as np
```

Consider the circuit where currents $X$ and $Y$ in two branches flow into a single node and then into a single $10\,\Omega$ resistor.


The currents in $X$ and $Y$ are independent random variables and their marginal probability density functions are given as:
+ The current $X$ is a Gaussian random variable with mean 3A and the standard deviation of 0.5A.
+ The current $Y$ has the marginal probability density function
$$
f_Y(y)=
    \begin{cases}
      0.5 & 0\,\text{A} \leq y \leq 2\,\text{A}\\
      0 & \text{otherwise}
    \end{cases}
$$


Find the probability that the current $X$ is larger than 3.83 A.

Represent your answer as `p = XX`.

<!-- grade_id: cell-ec2faf494a0e069b, points: 0, grade: false, solution: true -->
```python
p = 0.0485
```

<!-- grade_id: cell-e326bd77e75571b3, points: 1, grade: true, solution: false -->
```python
assert isinstance(p, float)

```

Find the second moment of $X$, $E[X^2]$.

Represent your answer as `ans = XX`.

<!-- grade_id: cell-7b060be93ba7990b, points: 0, grade: false, solution: true -->
```python
ans = 9.25
```

<!-- grade_id: cell-ed945f320ec84d1b, points: 1, grade: true, solution: false -->
```python
assert isinstance(ans, float)

```

Find the second moment of $Y$, $E[Y^2]$.

Represent your answer as `ans = XX`.

<!-- grade_id: cell-49a5d89c1efae188, points: 0, grade: false, solution: true -->
```python
ans = 4/3
```

<!-- grade_id: cell-cf535f1047779c92, points: 1, grade: true, solution: false -->
```python
assert isinstance(ans, float)

```

Find the correlation of $X$ and $Y$, $E[XY]$.

Represent your answer as `ans = XX`.

<!-- grade_id: cell-c20c2657b124f795, points: 0, grade: false, solution: true -->
```python
ans = 3.0
```

<!-- grade_id: cell-000f8779786513c8, points: 1, grade: true, solution: false -->
```python
assert isinstance(ans, float)

```

The sum of two currents $X+Y$ flow into a $10\,\Omega$ resistor.  The power dissipated over the resistor is then $10(X+Y)^2$.  Find the average power dissipated over the resistor.

<!-- grade_id: cell-edc5f414b2194293, points: 0, grade: false, solution: true -->
```python
ans = 165.83
```

<!-- grade_id: cell-d25097fa02e18ff3, points: 1, grade: true, solution: false -->
```python
assert isinstance(ans, float)

```