Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

# Electrons

```python
import numpy as np
```

An atom has two electrons. Each of the electrons can occupy one of the three orbits of this atom. Let the random variables $X$ and $Y$ denote the orbit number occupied by these two electrons. The joint probability distribution is given as

$f_{XY}(x,y)$ | $X=1$ | $X=2$ | $X=3$
--------------|-------|-------|-------
$Y=1$         | 0     |  1/8  | 0
$Y=2$         | 1/8   |  1/2  | 1/8
$Y=3$         | 0     |  1/8  | 0


Find the marginal pmf $p_X(x)$.

Represent your answer as a function `p_X(x)`.

<!-- grade_id: cell-f130733abe6c9daf, points: 0, grade: false, solution: true -->
```python
def p_X(x):
    if x == 1:
        return 1/8
    elif x == 2:
        return 3/4
    elif x == 3:
        return 1/8
    else:
        return 0.0
```

<!-- grade_id: cell-b30d501f8872800b, points: 2, grade: true, solution: false -->
```python
assert isinstance(p_X(1), float)

```

Find the marginal pmf $p_Y(y)$.

Represent your answer as a function `p_Y(y)`.

<!-- grade_id: cell-962fee63f3922fe6, points: 0, grade: false, solution: true -->
```python
def p_Y(y):
    if y == 1:
        return 1/8
    elif y == 2:
        return 3/4
    elif y == 3:
        return 1/8
    else:
        return 0.0
```

<!-- grade_id: cell-a56db365769053ee, points: 2, grade: true, solution: false -->
```python
assert isinstance(p_Y(1), float)

```

True or false:  $X$ and $Y$ are independent.

Respond with `ans = True` or `ans = False`.

<!-- grade_id: cell-fbb28642b26691e3, points: 0, grade: false, solution: true -->
```python
ans = False
```

<!-- grade_id: cell-b0da76172fbc9b41, points: 1, grade: true, solution: false -->
```python
assert isinstance(ans, bool)

```

Find $[X\leq 2|Y=2]$. 

Represent your answer as `p = XX`.

<!-- grade_id: cell-c6fa1517fd790cef, points: 0, grade: false, solution: true -->
```python
p = 5/6
```

<!-- grade_id: cell-8fd45902c98776c9, points: 1, grade: true, solution: false -->
```python
assert isinstance(p, float)

```