Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

# m-ary communications systems

```python
import numpy as np
```

In the course so far we have focused on binary communications systems where we are interested in transmitting one of two signals $-1$ or a $+1$.  However, in some communication systems we might be interested in sending $m$ different signals.  This is called $m$-ary communications systems (in \emph{bi}nary communications systems $m=2$).  Let's consider a hypothetical system where $X\in\{-2,-1,+1,+2\}$ and our receiver will decide between four different values $Y\in\{-2,-1,+1,+2\}$.  We can consider $X$ and $Y$ as joint random variables.  For example, we can interpret $P[X=-1,Y=+2]$ as the probability we tried transmitting a $X=-1$ and the receiver decided it was $Y=+2$.

| $p_{XY}(x,y)$ | $X=-2$ | $X=-1$ | $X=+1$ | $X=+2$ |
|---------------|--------|--------|--------|--------|
| $Y=-2$        | 0.18   | 0.03   | 0.01   | 0      |
| $Y=-1$        | 0.02   | 0.22   | 0.03   | 0.01   |
| $Y=+1$        | 0.01   | 0.03   | 0.22   | 0.02   |
| $Y=+2$        | 0      | 0.01   | 0.03   | 0.18   |

Find the marginal pmf $p_X(x)$.

Represent your answer as a dictionary `ans = {-2: XX, -1: XX, 1: XX, 2: XX}`.

<!-- grade_id: cell-6cad0ccd47c1d092, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a3bca194300d2db2, points: 1, grade: true, solution: false -->
```python
assert isinstance(ans, dict)

```

Find the marginal pmf $p_Y(y)$.

Represent your answer as a dictionary `ans = {-2: XX, -1: XX, 1: XX, 2: XX}`.

<!-- grade_id: cell-21614479d2e475cc, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e7bbdeb7a21f0cc3, points: 1, grade: true, solution: false -->
```python
assert isinstance(ans, dict)

```

True or False: these random variables are independent.

Represent your answer as `ans = True` or `ans = False`.

<!-- grade_id: cell-82c9286163003c73, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-30f9d243a96e3479, points: 1, grade: true, solution: false -->
```python
assert isinstance(ans, bool)

```

Find $P[X\leq 0 | Y = 2]$.

Represent your answer as `p = XX`.

<!-- grade_id: cell-eaf8bcf2cda7fa10, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2dcbc6bd14f60dda, points: 1, grade: true, solution: false -->
```python
assert isinstance(p, float)

```