Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

# Noise in communications

```python
import numpy as np
```

We often model a signal that is corrupted by noise as
$$
  Y = N + X.
$$
In this case, the noise $N$ is a Gaussian random variable with a zero mean and standard deviation of 0.5.  The signal $X$ is a binary signal where $P[X=+1]=0.4$ and $P[X=-1]=0.6$.

Find the first moment of $X$, $E[X]$.

Represent your answer as `ans = XX`.

<!-- grade_id: cell-1161fe476d2417bb, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1bccc0fb5d0dcf62, points: 1, grade: true, solution: false -->
```python
assert isinstance(ans, float)

```

Find the first moment of $N$, $E[N]$.

Represent your answer as `ans = XX`.

<!-- grade_id: cell-edda2cba49b8b88a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c2fef21baf16d151, points: 1, grade: true, solution: false -->
```python
assert isinstance(ans, float)

```

Assume that $X$ and $N$ are independent.  Find the correlation of $X$ and $N$, $E[XN]$.

Represent your answer as `ans = XX`.

<!-- grade_id: cell-d97f9e0552d21d19, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-60d6e3d8e39d838e, points: 1, grade: true, solution: false -->
```python
assert isinstance(ans, float)

```

We are concerned that the total signal $Y=X+N$ will exceed system limits such that $|Y| > 1.5$.  What is the probability that $P[|Y| > 1.5]$?

Represent your answer as `p = XX`.

<!-- grade_id: cell-baee58223adff7b9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7c90499c116c4de7, points: 2, grade: true, solution: false -->
```python
assert isinstance(p, float)

```